# Interact Exercise 2

## Imports

```python
%matplotlib inline
from matplotlib import pyplot as plt
import numpy as np
```

```python
from IPython.html.widgets import interact, interactive, fixed
from IPython.display import display
```

## Plotting with parameters

Write a `plot_sin1(a, b)` function that plots $sin(ax+b)$ over the interval $[0,4\pi]$.

* Customize your visualization to make it effective and beautiful.
* Customize the box, grid, spines and ticks to match the requirements of this data.
* Use enough points along the x-axis to get a smooth plot.
* For the x-axis tick locations use integer multiples of $\pi$.
* For the x-axis tick labels use multiples of pi using LaTeX: `$3\pi$`.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
plot_sine1(5, 3.4)
```

Then use `interact` to create a user interface for exploring your function:

* `a` should be a floating point slider over the interval $[0.0,5.0]$ with steps of $0.1$.
* `b` should be a floating point slider over the interval $[-5.0,5.0]$ with steps of $0.1$.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: interactex02a, points: 5, grade: true, solution: false -->
```python
assert True # leave this for grading the plot_sine1 exercise
```

In matplotlib, the line style and color can be set with a third argument to `plot`. Examples of this argument:

* dashed red: `r--`
* blue circles: `bo`
* dotted black: `k.`

Write a `plot_sine2(a, b, style)` function that has a third `style` argument that allows you to set the line style of the plot. The style should default to a blue line.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
plot_sine2(4.0, -1.0, 'r--')
```

Use `interact` to create a UI for `plot_sine2`.

* Use a slider for `a` and `b` as above.
* Use a drop down menu for selecting the line style between a **dotted blue line** line, **black circles** and **red triangles**.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: interactex02b, points: 5, grade: true, solution: false -->
```python
assert True # leave this for grading the plot_sine2 exercise
```