# Interact Exercise 01

## Import

```python
%matplotlib inline
from matplotlib import pyplot as plt
import numpy as np
```

```python
from IPython.html.widgets import interact, interactive, fixed
from IPython.display import display
```

## Interact basics

Write a `print_sum` function that `prints` the sum of its arguments `a` and `b`.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
def print_sum(a, b):
    """Print the sum of the arguments a and b."""
    # YOUR CODE HERE
    raise NotImplementedError()
```

Use the `interact` function to interact with the `print_sum` function.

* `a` should be a floating point slider over the interval `[-10., 10.]` with step sizes of `0.1`
* `b` should be an integer slider the interval [-8, 8] with step sizes of `2`.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: interactex01a, points: 5, grade: true, solution: false -->
```python
assert True # leave this for grading the print_sum exercise
```

Write a function named `print_string` that prints a string and additionally prints the length of that string if a boolean parameter is `True`.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
def print_string(s, length=False):
    """Print the string s and optionally its length."""
    # YOUR CODE HERE
    raise NotImplementedError()
```

Use the `interact` function to interact with the `print_string` function.

* `s` should be a textbox with the initial value `"Hello World!"`.
* `length` should be a checkbox with an initial value of `True`.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: interactex01b, points: 5, grade: true, solution: false -->
```python
assert True # leave this for grading the print_string exercise
```