# Project Euler: Problem 52

https://projecteuler.net/problem=52

It can be seen that the number, $125874$, and its double, $251748$, contain exactly the same digits, but in a different order.

Find the smallest positive integer, $x$, such that $2x$, $3x$, $4x$, $5x$, and $6x$, contain the same digits.

First, write a function `same_digits(x,y)` that returns `True` if two integers `x` and `y` have the exact same set of digits and multiplicities and `False` if they have different digits.

```python
%matplotlib inline
from matplotlib import pyplot as plt
import seaborn as sns
import numpy as np
```

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
def same_digits(x, y):
    """Do the integers x and y have the same digits, regardless of order."""
    # YOUR CODE HERE
    b = x.split()
    c = np.array(b)
    d = np.sort(c)
    e = int(d)
    f = y.split()
    g = np.array(f)
    h = np.sort(g)
    i = int(h)
    if i == e:
        return True
    else:
        return False
    
```

<!-- grade_id: projecteuler52a, points: 4, grade: true, solution: false -->
```python
assert same_digits('132', '321')
assert not same_digits('123', '3')
assert not same_digits('456', '0987654321')
```

Output:
```
<ipython-input-18-4d2369f19b4d> in <module>()
----> 1 assert same_digits('132', '321')
      2 assert not same_digits('123', '3')
      3 assert not same_digits('456', '0987654321')

<ipython-input-17-0260eab69cd7> in same_digits(x, y)
      6     c = np.array(b)
      7     d = np.sort(c)
----> 8     e = int(d)
      9     f = y.split()
     10     g = np.array(f)

TypeError: don't know how to convert scalar number to int
```

Now use the `same_digits` function to solve this Euler problem. As you work on this problem, be careful to debug and test your code on small integers before trying it on the full search.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
for x in range (0 , 1000):
    a = x
    b = 2 * x
    c = 3 * x
    d = 4 * x
    e = 5 * x
    f = 6 * x
    if len(str(a)) == 
            
 
```

<!-- grade_id: projecteuler52b, points: 6, grade: true, solution: false -->
```python
assert True # leave this cell to grade the solution
```