# Algorithms Exercise 1

## Imports

```python
%matplotlib inline
from matplotlib import pyplot as plt
import numpy as np
```

## Word counting

Write a function `tokenize` that takes a string of English text returns a list of words. It should also remove [stop words](http://en.wikipedia.org/wiki/Stop_words), which are common short words that are often removed before natural language processing. Your function should have the following logic:

* Split the string into lines using `splitlines`.
* Split each line into a list of words and merge the lists for each line.
* Use Python's builtin `filter` function to remove all punctuation.
* If `stop_words` is a list, remove all occurences of the words in the list.
* If `stop_words` is a space delimeted string of words, split them and remove them.
* Remove any remaining empty words.
* Make all words lowercase.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
def tokenize(s, stop_words= None, punctuation='`~!@#$%^&*()_-+={[}]|\:;"<,>.?/}\t'):
    """Split a string into a list of words, removing punctuation and stop words."""
    # YOUR CODE HERE
    #list comprehension
    u = list(filter(lambda x: x not in '~!@#$%^&*()_-+={[}]|\:;"<,>.?/}\t',s))
    d = ''.join(u)
    h = d.splitlines()
    q = []
    for i in q:
        j = r.lower()
        g = r.splt()
        q.append(g)
        w = sum(h, [])
    if type(stop_words) == str:
        sw = stop_words.splitlines()
        swa = []
        for j in sw:
            o = j.lower()
            cv = o.split()
            swa.append(cw)
            be = sum(swa , [])
        return be
    elif type(stop_words)==list:
        dr = stop_words
    for f in w:
        for l in dr:
            if l == f:
                w.remove(l)
    return w
```

<!-- grade_id: algorithmsex01a, points: 2, grade: true, solution: false -->
```python
assert tokenize("This, is the way; that things will end", stop_words=['the', 'is']) == \
    ['this', 'way', 'that', 'things', 'will', 'end']
wasteland = """
APRIL is the cruellest month, breeding
Lilacs out of the dead land, mixing
Memory and desire, stirring
Dull roots with spring rain.
"""

assert tokenize(wasteland, stop_words='is the of and') == \
    ['april','cruellest','month','breeding','lilacs','out','dead','land',
     'mixing','memory','desire','stirring','dull','roots','with','spring',
     'rain']
```

Output:
```
<ipython-input-61-cb7875c10e3e> in <module>()
----> 1 assert tokenize("This, is the way; that things will end", stop_words=['the', 'is']) ==     ['this', 'way', 'that', 'things', 'will', 'end']
      2 wasteland = """
      3 APRIL is the cruellest month, breeding
      4 Lilacs out of the dead land, mixing
      5 Memory and desire, stirring

<ipython-input-60-aba21a402eb4> in tokenize(s, stop_words, punctuation)
     23     elif type(stop_words)==list:
     24         dr = stop_words
---> 25     for f in w:
     26         for l in dr:
     27             if l == f:

UnboundLocalError: local variable 'w' referenced before assignment
```

Write a function `count_words` that takes a list of words and returns a dictionary where the keys in the dictionary are the unique words in the list and the values are the word counts.

<!-- grade_id: algorithmsex01b, points: 2, grade: false, solution: true -->
```python
def count_words(data):
    """Return a word count dictionary from the list of words in data."""
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: algorithmsex01b, points: 2, grade: true, solution: false -->
```python
assert count_words(tokenize('this and the this from and a a a')) == \
    {'a': 3, 'and': 2, 'from': 1, 'the': 1, 'this': 2}
```

Output:
```
<ipython-input-18-71b7cc9f406e> in <module>()
----> 1 assert count_words(tokenize('this and the this from and a a a')) ==     {'a': 3, 'and': 2, 'from': 1, 'the': 1, 'this': 2}

<ipython-input-15-1bd47f096613> in tokenize(s, stop_words, punctuation)
      3     # YOUR CODE HERE
      4     #list comprehension
----> 5     p = s.split[" "]
      6     return p

TypeError: 'builtin_function_or_method' object is not subscriptable
```

Write a function `sort_word_counts` that return a list of sorted word counts:

* Each element of the list should be a `(word, count)` tuple.
* The list should be sorted by the word counts, with the higest counts coming first.
* To perform this sort, look at using the `sorted` function with a custom `key` and `reverse`
  argument.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
def sort_word_counts(wc):
    """Return a list of 2-tuples of (word, count), sorted by count descending."""
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: algorithmsex01c, points: 2, grade: true, solution: false -->
```python
assert sort_word_counts(count_words(tokenize('this and a the this this and a a a'))) == \
    [('a', 4), ('this', 3), ('and', 2), ('the', 1)]
```

Output:
```
<ipython-input-20-2b19d16a049a> in <module>()
----> 1 assert sort_word_counts(count_words(tokenize('this and a the this this and a a a'))) ==     [('a', 4), ('this', 3), ('and', 2), ('the', 1)]

<ipython-input-15-1bd47f096613> in tokenize(s, stop_words, punctuation)
      3     # YOUR CODE HERE
      4     #list comprehension
----> 5     p = s.split[" "]
      6     return p

TypeError: 'builtin_function_or_method' object is not subscriptable
```

Perform a word count analysis on Chapter 1 of Moby Dick, whose text can be found in the file `mobydick_chapter1.txt`:

* Read the file into a string.
* Tokenize with stop words of `'the of and a to in is it that as'`.
* Perform a word count, the sort and save the result in a variable named `swc`.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
<ipython-input-21-a3eb14613b08> in <module>()
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: algorithmsex01d, points: 2, grade: true, solution: false -->
```python
assert swc[0]==('i',43)
assert len(swc)==848
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-22-79f30673d9f4> in <module>()
----> 1 assert swc[0]==('i',43)
      2 assert len(swc)==848

NameError: name 'swc' is not defined
```

Create a "Cleveland Style" [dotplot](http://en.wikipedia.org/wiki/Dot_plot_%28statistics%29) of the counts of the top 50 words using Matplotlib. If you don't know what a dotplot is, you will have to do some research...

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
<ipython-input-23-a3eb14613b08> in <module>()
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: algorithsex01e, points: 2, grade: true, solution: false -->
```python
assert True # use this for grading the dotplot
```