# Matplotlib Exercise 1

## Imports

```python
%matplotlib inline
import matplotlib.pyplot as plt
import numpy as np
```

## Line plot of sunspot data

Download the `.txt` data for the "Yearly mean total sunspot number [1700 - now]" from the [SILSO](http://www.sidc.be/silso/datafiles) website. Upload the file to the same directory as this notebook.

<!-- grade_id: matplotlibex01a, points: 1, grade: true, solution: false -->
```python
import os
assert os.path.isfile('yearssn.dat')
```

Use `np.loadtxt` to read the data into a NumPy array called `data`. Then create two new 1d NumPy arrays named `years` and `ssc` that have the sequence of year and sunspot counts.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: matplotlibex01b, points: 1, grade: true, solution: false -->
```python
assert len(year)==315
assert year.dtype==np.dtype(float)
assert len(ssc)==315
assert ssc.dtype==np.dtype(float)
```

Make a line plot showing the sunspot count as a function of year.

* Customize your plot to follow Tufte's principles of visualizations.
* Adjust the aspect ratio/size so that the steepest slope in your plot is *approximately* 1.
* Customize the box, grid, spines and ticks to match the requirements of this data.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: matplotlibex01c, points: 3, grade: true, solution: false -->
```python
assert True # leave for grading
```

Describe the choices you have made in building this visualization and how they make it effective.

<!-- grade_id: matplotlibex01d, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

Now make 4 subplots, one for each century in the data set. This approach works well for this dataset as it allows you to maintain mild slopes while limiting the overall width of the visualization. Perform similar customizations as above:

* Customize your plot to follow Tufte's principles of visualizations.
* Adjust the aspect ratio/size so that the steepest slope in your plot is *approximately* 1.
* Customize the box, grid, spines and ticks to match the requirements of this data.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: matplotlibex01e, points: 4, grade: true, solution: false -->
```python
assert True # leave for grading
```