# Integration Exercise 1

## Imports

```python
%matplotlib inline
import matplotlib.pyplot as plt
import numpy as np
from scipy import integrate
```

## Trapezoidal rule

The [trapezoidal](http://en.wikipedia.org/wiki/Trapezoidal_rule) rule generates a numerical approximation to the 1d integral:

$$ I(a,b) = \int_a^b f(x) dx $$

by dividing the interval $[a,b]$ into $N$ subdivisions of length $h$:

$$ h = (b-a)/N $$

Note that this means the function will be evaluated at $N+1$ points on $[a,b]$. The main idea of the trapezoidal rule is that the function is approximated by a straight line between each of these points.

Write a function `trapz(f, a, b, N)` that performs trapezoidal rule on the function `f` over the interval $[a,b]$ with `N` subdivisions (`N+1` points).

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
def trapz(f, a, b, N):
    """Integrate the function f(x) over the range [a,b] with N points."""
    # YOUR CODE HERE
    h = (b-a)/N
    n = np.arange(1,N)
    I = h*(0.5*f(a) + 0.5*f(b) + (f(a+n*h)).sum())
    return I
```

```python
f = lambda x: x**2
g = lambda x: np.sin(x)
```

<!-- grade_id: integrationex01a, points: 5, grade: true, solution: false -->
```python
I = trapz(f, 0, 1, 1000)
assert np.allclose(I, 0.33333349999999995)
J = trapz(g, 0, np.pi, 1000)
assert np.allclose(J, 1.9999983550656628)
```

Now use `scipy.integrate.quad` to integrate the `f` and `g` functions and see how the result compares with your `trapz` function. Print the results and errors.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
fnew, err1 = integrate.quad(f, 0.0, 1)
gnew, err2 = integrate.quad(g, 0.0, np.pi)
print(I, fnew, err1)
print(J, gnew, err1)

```

Output:
```
0.3333335 0.33333333333333337 3.700743415417189e-15
1.99999835507 2.0 3.700743415417189e-15
```

<!-- grade_id: integrationex01b, points: 5, grade: true, solution: false -->
```python
assert True # leave this cell to grade the previous one
```