# Integration Exercise 2

## Imports

```python
%matplotlib inline
import matplotlib.pyplot as plt
import numpy as np
import seaborn as sns
from scipy import integrate
import math as m
```

## Indefinite integrals

Here is a [table of definite integrals](http://en.wikipedia.org/wiki/List_of_definite_integrals). Many of these integrals has a number of parameters $a$, $b$, etc.

Find five of these integrals and perform the following steps:

1. Typeset the integral using LateX in a Markdown cell.
1. Define an `integrand` function that computes the value of the integrand.
2. Define an `integral_approx` funciton that uses `scipy.integrate.quad` to peform the integral.
3. Define an `integral_exact` function that computes the exact value of the integral.
4. Call and print the return value of `integral_approx` and `integral_exact` for one set of parameters.

Here is an example to show what your solutions should look like:

### Example

Here is the integral I am performing:

$$ I_1 = \int_0^\infty \frac{dx}{x^2 + a^2} = \frac{\pi}{2a} $$

```python
def integrand(x, a):
    return 1.0/(x**2 + a**2)

def integral_approx(a):
    # Use the args keyword argument to feed extra arguments to your integrand
    I, e = integrate.quad(integrand, 0, np.inf, args=(a,))
    return I

def integral_exact(a):
    return 0.5*np.pi/a

print("Numerical: ", integral_approx(1.0))
print("Exact    : ", integral_exact(1.0))

```

Output:
```
Numerical:  1.5707963267948966
Exact    :  1.5707963267948966
```

```python
assert True # leave this cell to grade the above integral
```

### Integral 1

<!-- grade_id: integrationex03a, points: 1, grade: true, solution: true -->
Here is the integral I am performing:

$$ I_1 =  \int_0^\frac{\pi}{2} sin^2 x dx = \frac{\pi}{4}  $$

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
# raise NotImplementedError()
def integrand(x):
    return (np.sin(x))**2

def integral_approx(a):
    I, err = integrate.quad(integrand, 0, np.pi)
    return I

def integral_exact(a):
    return np.pi/4

print("Numerical: ", integral_approx(1.0))
print("Exact    : ", integral_exact(1.0))
```

Output:
```
Numerical:  1.5707963267948966
Exact    :  0.7853981633974483
```

<!-- grade_id: integrationex03b, points: 1, grade: true, solution: false -->
```python
assert True # leave this cell to grade the above integral
```

### Integral 2

<!-- grade_id: integrationex03c, points: 1, grade: true, solution: true -->
$$ I_2 = \int_0^\infty \frac{sin^2 px}{x}dx = \frac{\pi p}{2} $$

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def integrand(x, p):
    return ((np.sin(p*x))**2)/x

def integral_approx(p):
    # Use the args keyword argument to feed extra arguments to your integrand
    I, e = integrate.quad(integrand, 0, np.inf, args=(p,))
    return I

def integral_exact(p):
    return 0.5*np.pi*p/2

print("Numerical: ", integral_approx(1.0))
print("Exact    : ", integral_exact(1.0))
```

Output:
```
Numerical:  5.066861106505159
Exact    :  0.7853981633974483
```

<!-- grade_id: integrationex03d, points: 1, grade: true, solution: false -->
```python
assert True # leave this cell to grade the above integral
```

### Integral 3

<!-- grade_id: integrationex03e, points: 1, grade: true, solution: true -->
$$ I_3 = \int_0^\infty \frac{1 - cos px}{x}dx = \frac{\pi p}{2} $$

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def integrand(x, p):
    return (1.0-np.cos(p*x))/x

def integral_approx(p):
    # Use the args keyword argument to feed extra arguments to your integrand
    I, e = integrate.quad(integrand, 0, np.inf, args=(p,))
    return I

def integral_exact(p):
    return 0.5*np.pi*p/2

print("Numerical: ", integral_approx(1.0))
print("Exact    : ", integral_exact(1.0))
```

Output:
```
Numerical:  7.460786724542071
Exact    :  0.7853981633974483
```

<!-- grade_id: integrationex03f, points: 1, grade: true, solution: false -->
```python
assert True # leave this cell to grade the above integral
```

### Integral 4

<!-- grade_id: integrationex03g, points: 1, grade: true, solution: true -->
$$ I_4 = \int_0^\infty \frac{x}{e^x-1}  dx = \frac{\pi^2}{6} $$

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def integrand(x):
    return x/(np.exp(x)-1)

def integral_approx(a):
    I, err = integrate.quad(integrand, 0, np.inf)
    return I

def integral_exact(a):
    return np.pi**2/6

print("Numerical: ", integral_approx(1.0))
print("Exact    : ", integral_exact(1.0))
```

Output:
```
Numerical:  1.6449340668482264
Exact    :  1.6449340668482264
```

<!-- grade_id: integrationex03h, points: 1, grade: true, solution: false -->
```python
assert True # leave this cell to grade the above integral
```

### Integral 5

<!-- grade_id: integrationex03i, points: 1, grade: true, solution: true -->
$$ I_5 = \int_0^\infty \frac{x}{e^x+1}  dx = \frac{\pi^2}{12} $$

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
# YOUR CODE HERE
def integrand(x):
    return x/(np.exp(x)+1)

def integral_approx(a):
    I, err = integrate.quad(integrand, 0, np.inf)
    return I

def integral_exact(a):
    return np.pi**2/12

print("Numerical: ", integral_approx(1.0))
print("Exact    : ", integral_exact(1.0))
```

Output:
```
Numerical:  0.822467033424113
Exact    :  0.8224670334241132
```

<!-- grade_id: integrationex03j, points: 1, grade: true, solution: false -->
```python
assert True # leave this cell to grade the above integral
```