# Numpy Exercise 2

## Imports

```python
import numpy as np
%matplotlib inline
import matplotlib.pyplot as plt
import seaborn as sns
```

## Factorial

Write a function that computes the factorial of small numbers using `np.arange` and `np.cumprod`.

```python
def np_fact(n):
    a = np.arange(1,n)
    n = n * (n+1)
    return n
print np_fact(8)
    
```

Output:
```
72
```

<!-- grade_id: numpyex02a, points: 4, grade: true, solution: false -->
```python
assert np_fact(0)==1
assert np_fact(1)==1
assert np_fact(10)==3628800
assert [np_fact(i) for i in range(0,11)]==[1,1,2,6,24,120,720,5040,40320,362880,3628800]
```

Output:
```
<ipython-input-13-c7ed3614e378> in <module>()
----> 1 assert np_fact(0)==1
      2 assert np_fact(1)==1
      3 assert np_fact(10)==3628800
      4 assert [np_fact(i) for i in range(0,11)]==[1,1,2,6,24,120,720,5040,40320,362880,3628800]

<ipython-input-12-0cdb34430db8> in np_fact(n)
      1 def np_fact(n):
      2     a = np.arange(1,n)
----> 3     b = np.factorials(a)
      4     return b
      5 

AttributeError: 'module' object has no attribute 'factorials'
```

Write a function that computes the factorial of small numbers using a Python loop.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
def loop_fact(n):
    q = 1
    for i in range(n):
        q = q * (i+1)
    return q

print np_fact(5)
```

Output:
```
120
```

<!-- grade_id: numpyex02b, points: 4, grade: true, solution: false -->
```python
assert loop_fact(0)==1
assert loop_fact(1)==1
assert loop_fact(10)==3628800
assert [loop_fact(i) for i in range(0,11)]==[1,1,2,6,24,120,720,5040,40320,362880,3628800]
```

Use the `%timeit` magic to time both versions of this function for an argument of `50`. The syntax for `%timeit` is:

```python
%timeit -n1 -r1 function_to_time()
```

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

In the cell below, summarize your timing tests. Which version is faster? Why do you think that version is faster?

<!-- grade_id: numpyex02c, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE