# Numpy Exercise 2

## Imports

```python
import numpy as np
%matplotlib inline
import matplotlib.pyplot as plt
import seaborn as sns
```

## Factorial

Write a function that computes the factorial of small numbers using `np.arange` and `np.cumprod`.

```python
def np_fact(n):
    if n == 0:
        return 1
    elif n >0:
        a = np.arange(1,n,1)
        n_1 = a * (n+1)
    return n_1
print np_fact(10)
    
```

Output:
```
110
```

<!-- grade_id: numpyex02a, points: 4, grade: true, solution: false -->
```python
assert np_fact(0)==1
assert np_fact(1)==1
assert np_fact(10)==3628800
assert [np_fact(i) for i in range(0,11)]==[1,1,2,6,24,120,720,5040,40320,362880,3628800]
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-7-c7ed3614e378> in <module>()
      1 assert np_fact(0)==1
----> 2 assert np_fact(1)==1
      3 assert np_fact(10)==3628800
      4 assert [np_fact(i) for i in range(0,11)]==[1,1,2,6,24,120,720,5040,40320,362880,3628800]

AssertionError:
```

Write a function that computes the factorial of small numbers using a Python loop.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
def loop_fact(n):
    q = 1
    for i in range(n):
        q = q * (i+1)
    return q

print np_fact(5)
```

Output:
```
30
```

<!-- grade_id: numpyex02b, points: 4, grade: true, solution: false -->
```python
assert loop_fact(0)==1
assert loop_fact(1)==1
assert loop_fact(10)==3628800
assert [loop_fact(i) for i in range(0,11)]==[1,1,2,6,24,120,720,5040,40320,362880,3628800]
```

Use the `%timeit` magic to time both versions of this function for an argument of `50`. The syntax for `%timeit` is:

```python
%timeit -n1 -r1 function_to_time()
```

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
<ipython-input-10-a3eb14613b08> in <module>()
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```

In the cell below, summarize your timing tests. Which version is faster? Why do you think that version is faster?

<!-- grade_id: numpyex02c, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE