# Numpy Exercise 3

## Imports

```python
import numpy as np
%matplotlib inline
import matplotlib.pyplot as plt
import seaborn as sns
```

```python
import antipackage
import github.ellisonbg.misc.vizarray as va
```

## Geometric Brownian motion

Here is a function that produces standard Brownian motion using NumPy. This is also known as a [Wiener Process](http://en.wikipedia.org/wiki/Wiener_process).

```python
def brownian(maxt, n):
    """Return one realization of a Brownian (Wiener) process with n steps and a max time of t."""
    t = np.linspace(0.0,maxt,n)
    h = t[1]-t[0]
    Z = np.random.normal(0.0,1.0,n-1)
    dW = np.sqrt(h)*Z
    W = np.zeros(n)
    W[1:] = dW.cumsum()
    return t, W
```

Call the `brownian` function to simulate a Wiener process with `1000` steps and max time of `1.0`. Save the results as two arrays `t` and `W`.

<!-- grade_id: numpyex03a, points: 2, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: numpyex03a, points: 2, grade: true, solution: false -->
```python
assert isinstance(t, np.ndarray)
assert isinstance(W, np.ndarray)
assert t.dtype==np.dtype(float)
assert W.dtype==np.dtype(float)
assert len(t)==len(W)==1000
```

Visualize the process using `plt.plot` with `t` on the x-axis and `W(t)` on the y-axis. Label your x and y axes.

<!-- grade_id: numpyex03b, points: 2, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: numpyex03b, points: 2, grade: true, solution: false -->
```python
assert True # this is for grading
```

Use `np.diff` to compute the changes at each step of the motion, `dW`, and then compute the mean and standard deviation of those differences.

<!-- grade_id: numpyex03c, points: 2, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: numpyex03c, points: 2, grade: true, solution: false -->
```python
assert len(dW)==len(W)-1
assert dW.dtype==np.dtype(float)
```

Write a function that takes $W(t)$ and converts it to geometric Brownian motion using the equation:

$$
X(t) = X_0 e^{((\mu - \sigma^2/2)t + \sigma W(t))}
$$

Use Numpy ufuncs and no loops in your function.

<!-- grade_id: numpyex03d, points: 2, grade: false, solution: true -->
```python
def geo_brownian(t, W, X0, mu, sigma):
    "Return X(t) for geometric brownian motion with drift mu, volatility sigma."""
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: numpyex03d, points: 2, grade: true, solution: false -->
```python
assert True # leave this for grading
```

Use your function to simulate geometric brownian motion, $X(t)$ for $X_0=1.0$, $\mu=0.5$ and $\sigma=0.3$ with the Wiener process you computed above.

Visualize the process using `plt.plot` with `t` on the x-axis and `X(t)` on the y-axis. Label your x and y axes.

<!-- grade_id: numpyex03f, points: 2, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: numpyex03e, points: 2, grade: true, solution: false -->
```python
assert True # leave this for grading
```