# Numpy Exercise 1

## Imports

```python
import numpy as np
%matplotlib inline
import matplotlib.pyplot as plt
import seaborn as sns
```

```python
import antipackage
import github.ellisonbg.misc.vizarray as va
```

Output:
```
Downloading:  https://raw.githubusercontent.com/ellisonbg/misc/master/vizarray.py
Using existing version:  github.ellisonbg.misc.vizarray
```

## Checkerboard

Write a Python function that creates a square `(size,size)` 2d Numpy array with the values `0.0` and `1.0`:

* Your function should work for both odd and even `size`.
* The `0,0` element should be `1.0`.
* The `dtype` should be `float`.

```python
va.enable()
```

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
def checkerboard(size):
    k = np.zeros((size,size))
    for i in range(size):
        if size %2 == 0:
            k[i,i%2:size:2] = 1.0
        elif size %2 == 1:
            k[i,i%2:size+1:2] = 1.0
    return k

print checkerboard(4)

```

Output:
```
[[ 1.  0.  1.  0.]
 [ 0.  1.  0.  1.]
 [ 1.  0.  1.  0.]
 [ 0.  1.  0.  1.]]
```

```python
va.vizarray(checkerboard(5)) 
```

Output:
```
<ipythonblocks.ipythonblocks.BlockGrid at 0x7f1f6c1d9bd0>
```

<!-- grade_id: numpyex01a, points: 6, grade: true, solution: false -->
```python
a = checkerboard(4)
assert a[0,0]==1.0
assert a.sum()==8.0
assert a.dtype==np.dtype(float)
assert np.all(a[0,0:5:2]==1.0)
assert np.all(a[1,0:5:2]==0.0)

b = checkerboard(5)
assert b[0,0]==1.0
assert b.sum()==13.0
assert np.all(b.ravel()[0:26:2]==1.0)
assert np.all(b.ravel()[1:25:2]==0.0)
```

Use `vizarray` to visualize a checkerboard of `size=20` with a block size of `10px`.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
# raise NotImplementedError()
m = checkerboard(20)
va.set_block_size(10)
va.vizarray(m)
```

Output:
```
<ipythonblocks.ipythonblocks.BlockGrid at 0x7f1f3d59e1d0>
```

<!-- grade_id: numpyex01b, points: 2, grade: true, solution: false -->
```python
assert True
```

Use `vizarray` to visualize a checkerboard of `size=27` with a block size of `5px`.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
# raise NotImplementedError()
r = checkerboard(27)
va.set_block_size(5)
va.vizarray(r)
```

Output:
```
<ipythonblocks.ipythonblocks.BlockGrid at 0x7f1f3d5c5f50>
```

<!-- grade_id: numpyex01c, points: 2, grade: true, solution: false -->
```python
assert True
```

```python
def checkerboard(w, h) :
    re = np.r_[ w*[0,1] ]
    ro = np.r_[ w*[1,0] ]
    return np.row_stack(h*(re, ro))


AB = checkerboard(5, 10)
```

```python
va.enable()
```

```python
def gen_checkerboard(size):
    even = 1 * [0,1]
    odd = 1* [1,0]
    checkerboard = np.row_stack(size/2*(even, odd))
    return checkerboard.repeat(size, axis = 0)
print (gen_checkerboard(2))
```

Output:
```
[[0 1]
 [0 1]
 [1 0]
 [1 0]]
```

```python

```